---
title       : Navigation in R
description : Basic navigation in R, common functions, indexing, help, etc.



--- type:NormalExercise lang:r xp:100 skills:1 key:17b7ba7439
## <<<New Exercise>>>



*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
